---
title       : Advanced testwhat functions
description : Some questions about advanced functions in the testwhat package.

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:1c2b6es822
## test_data_frame

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:1c2b6es823
## test_if_else

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:1c2b6es824
## test_for_loop

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:1c2b6es825
## test_while_loop

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:1c2b6es826
## test_function_definition

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```