---
title       : Logic-inducing testwhat functions
description : Some questions about logic-inducing functions in the testwhat package.


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:2cec9c954b
## test_correct (1)

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:54b2cec9c9
## test_correct (2)

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:9c154c2cec
## test_correct (3)

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:2c053bec9c
## test_or

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```

