---
title       : testwhat case studies
description : Some case studies to see how different testwhat functions can work together

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:cb2ce954c9
## Case 1

Case studies will be added to this chapter as time goes and requests for clarifications pour in.

*** =instructions
- OK!

*** =hint
No hints, I'm sorry!

*** =sct
```{r}
msg1 <- "Glad we're on the same page!"
test_mc(1, feedback_msgs = msg1)
```