---
title       : testwhat case studies
description : Some case studies to see how different testwhat functions can work together

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:cb2ce954c9
## Case 1

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
No hints, I'm sorry!

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```